```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment2.ipynb")
```

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Probability, Statistics and Distances Homework

Please complete the following assignment. Run all cells and submit your completed notebook through Gradescope.

This assignment covers the following topics:

- Frequentist vs Bayesian probability
- Definition of probability
- Conditional probability and Bayes theorem
- Random variables
- Expectation, variance, and covariance
- Binomial distribution
- Poisson distribution
- Gaussian distribution
- Jaccard Index
- Euclidean Distances

<!-- BEGIN QUESTION -->

#### 1. Which of the following statements best describes the Bayesian interpretation of probability?

    A) Probability is the long-run relative frequency of an event occurring in repeated trials.
    B) Probability represents a degree of belief or confidence in the occurrence of an event.
    C) Probability is the ratio of favorable outcomes to the total number of equally likely outcomes.
    D) Probability is the measure of the likelihood that a random variable takes a specific value.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### 2. Write a function `time_to_six() -> tuple` that simulates rolling a fair six-sided die repeatedly until a `6` appears.  
The function should return a tuple `(rolls, count)`  

```python
import numpy as np

def time_to_six() -> tuple:
    roll = []
    while True:
        a = np.random.randint(1, 7)
        roll.append(a)
        if a == 6:
            break
    return (roll,len(roll) )
    
    ...
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

#### Question 3. 

Suppose that 0.5% of all credit card transactions are fraudulent. A fraud detection system correctly flags fraudulent transactions 96% of the time (sensitivity), and incorrectly flags legitimate transactions 3% of the time (false positive rate). If a randomly selected transaction is flagged, what is the probability that it is actually fraudulent?

    Return your answer as a float between 0 and 1.

```python
def answer3():
    P_F = 0.005         
    P_not_F = 0.995     
    P_Flag_given_F = 0.96      
    P_Flag_given_not_F = 0.03  
    # Total probability that a transaction is flagged
    P_Flag = P_Flag_given_F * P_F + P_Flag_given_not_F * P_not_F
    
    P_F_given_Flag = (P_Flag_given_F * P_F) / P_Flag
    
    return P_F_given_Flag
    
    
    ...
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

#### Question 4. 

Let X be a discrete random variable that takes the values 1, 3, 5 with probabilities P(X=1)=0.4, P(X=3)=0.2, P(X=5)=0.4.

- Compute the expected value E[X].  
- Return your answer as a float.  

```python
def ans4():
    return 1*0.4 + 3*0.2 + 5*0.4
    ...
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

####  Question 5. 

Let X be a random variable as in Question 4. Compute the variance Var(X).
-    Return your answer as a float.

```python
def ans5():
    E_X = ans4() 
    E_X2 = 1**2*0.4 + 3**2*0.2 + 5**2*0.4
    var_X = E_X2 - E_X**2
    return var_X  
    ...

```

```python
ans5()
```

Output:
```
3.1999999999999993
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

#### Question 6. 

A fair coin is tossed 10 times. What is the probability of getting exactly 6 heads?

    Return your answer as a float.

```python
from math import comb
def ans6():
    n = 10
    p = .5
    k = 6
    prob = comb(n,k) * (p**k) * ((1-p)**(n-k))
    return prob
    ...
```

```python
ans6()
```

Output:
```
0.205078125
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

#### Question 7. 

A radioactive source emits particles following a Poisson distribution with an average rate of λ = 4 particles per minute. 

    What is the probability that exactly 6 particles are emitted in a given minute?

    Return your answer as a float.

```python
from scipy.stats import poisson
def ans7():
    λ = 4
    k = 6
    prob = poisson.pmf(k, λ)
    return prob 
    ...
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

#### Question 8. 

Assuming that the heights of adult males are normally distributed with a mean of 70 inches and a standard deviation of 3 inches, what is the probability that a randomly selected male is taller than 75 inches?

    Return your answer as a float.

```python
from scipy.stats import norm
def ans8():
    mu = 70
    sigma = 3
    x = 75
    prob = 1 - norm.cdf(x, loc=mu, scale=sigma)
    return prob
    ...

ans8()
```

Output:
```
np.float64(0.047790352272814696)
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2 (points: 1)
```
>>> (rolls, count) = time_to_six()
>>> assert isinstance(rolls, list)
>>> assert isinstance(count, int)
>>> assert all((1 <= x <= 6 for x in rolls))
>>> assert rolls[-1] == 6
>>> assert count == len(rolls)
```


Tests for q3 (points: 1)
```
>>> assert answer3() > 0.13 and answer3() < 0.15
>>> assert type(answer3()) == float
```


Tests for q4 (points: 1)
```
>>> assert type(ans4()) == float
>>> assert 2.9 < ans4() < 3.1
```


Tests for q5 (points: 1)
```
>>> assert type(ans5()) == float
>>> assert abs(ans5() - 3.2) < 1e-09
>>> assert 3.19 < ans5() < 3.21
```


Tests for q6 (points: 1)
```
>>> assert type(ans6()) == float
>>> assert 0.2 < ans6() < 0.21
>>> assert abs(ans6()) - 0.2 < 0.05
```


Tests for q7 (points: 1)
```
>>> def test_ans7(ans7):
...     assert type(ans7()) == float
...     assert 0.1 < ans7() < 0.11
...     assert abs(ans7()) - 0.1 < 0.0042
```


Tests for q8 (points: 1)
```
>>> def test_ans8(ans8):
...     assert type(ans8()) == float
...     assert 0.04 < abs(ans8()) < 0.05
```
